# Gradient Descent - Solutions

### Setup and Preparation

In the following, we want to find the minimum of function ``f_quad`` using gradient descent.
The functions ``f_quad_dw`` and ``f_quad_db`` are the partial derivatives of ``f_quad``.

Make sure to execute the following code cells so that the definitions are loaded into your kernel.

```python
def f_quad(w, b):
    return (w * 2.5 - 3)**2 + (b * 1.7 + 1)**2

def f_quad_dw(w, b):
    return 2 * (w * 2.5 - 3) * 2.5

def f_quad_db(w, b):
    return 2 * (b * 1.7 + 1) * 1.7
```

```python
import numpy as np

def f_log(w, b):
    return np.log(f_quad(w, b))
```

The following code cells create plots of the function ``f_quad``.

Make sure to execute the following code cells to have the definitions loaded to your kernel.

```python
import matplotlib.pyplot as plt
import numpy as np
from matplotlib import cm

def get_value_grid(x_min, x_max, y_min, y_max, z_func):
    xs = np.linspace(x_min, x_max, num=50, endpoint=True)
    ys = np.linspace(y_min, y_max, num=50, endpoint=True)
    X, Y = np.meshgrid(xs, ys)
    Z = np.zeros_like(X)
    for i in range(X.shape[0]):
        Z[i, :] = z_func(X[i, :], Y[i, :])
    return X, Y, Z

def plot_3D(X, Y, Z):
    fig, ax = plt.subplots(subplot_kw={"projection": "3d"})
    # Plot the surface.
    surf = ax.plot_surface(X, Y, Z, 
                           cmap=cm.rainbow,
                           linewidth=0, 
                           antialiased=False)
            
X, Y, Z = get_value_grid(-10, 10, -10, 10, f_quad)
plot_3D(X, Y, Z)
plt.show()
```

```python
def plot_2D_color_encoded(X, Y, Z):    
    plt.figure()
    plt.scatter(X, Y, c=Z, cmap=cm.rainbow)
    plt.colorbar()
    plt.xlabel("X")
    plt.ylabel("Y")

plt.figure()
plot_2D_color_encoded(X, Y, Z)
plt.show()
```

## Exercise: Gradient Descent

Implement the function ``gradient_descent_step_2d(x_0, y_0, target_fn, target_fn_dx, target_fn_dy, step_size)``
which performs one step of gradient descent for a function ``target_fn`` which takes two arguments ``x,y``.
The function takes the current position ``x_0, y_0``,
the function to be minimized,
its partial derivatives w.r.t. ``x`` and ``y``,
as well as the step size for the gradient descent step as input.
A tuple ``x_1, y_1`` indicating the position after one step of gradient descent shall be returned.

<!-- grade_id: cell-0864c400fda9, points: 0, grade: false, solution: true -->
```python
def gradient_descent_step_2d(x_0, y_0, target_fn, target_fn_dx, target_fn_dy, step_size):
    # start solution
    d1 = target_fn_dx(x_0, y_0)
    d2 = target_fn_dy(x_0, y_0)
    x_1 = x_0 - step_size * d1
    y_1 = y_0 - step_size * d2
    return x_1, y_1
    # end solution
```

The following code cell contains a small test case for your implementation.

You should be able to execute it without any errors.

<!-- grade_id: cell-4741fd048fc7, points: 1, grade: true, solution: false -->
```python
x_1, y_1 = gradient_descent_step_2d(0, 1, lambda x, y: 2*x - y, lambda x, y: 2, lambda x, y: -1, 0.1)
assert np.isclose(x_1, -0.2)
assert np.isclose(y_1, 1.1)
```

Implement a function ``gradient_descent(x_0, y_0, target_fn, target_fn_dx, target_fn_dy, step_size, stop_eps, max_iter)``
which performs multiple steps of gradient descent for a real-valued target function ``target_fn`` with two inputs ``x, y``.
Write the position after each iteration of gradient descent in two arrays ``xs`` and ``ys``.
The gradient descent stops, when the squared Euclidean distance between two succeeding positions is smaller than some threshold value ``stop_eps``,
or after performing ``max_iter`` steps, whichever comes first.

The function returns the two arrays ``xs`` and ``ys`` containing the positions after each iteration of gradient descent,
including the starting positions ``x_0`` and ``y_0``, respectively. 

Within your function, consider using print statements for debugging and tracing purposes.

<!-- grade_id: cell-9ea2c1231799, points: 1, grade: true, solution: true -->
```python
def gradient_descent(x_0, y_0, target_fn, target_fn_dx, target_fn_dy, step_size, stop_eps, max_iter):
    xs = [x_0]
    ys = [y_0]
    f_value = target_fn(xs[-1], ys[-1])
    # start solution
    print("{:20s} - {:>14s}, {:>24s}".format("Iteration", "Step Size", "Target function value"))
    for i in range(max_iter):
        x, y = gradient_descent_step_2d(xs[-1], ys[-1], target_fn, target_fn_dx, target_fn_dy, step_size)
        diff = (x - xs[-1]) ** 2 + (y - ys[-1]) ** 2
        xs.append(x)
        ys.append(y)
        f_value = target_fn(xs[-1], ys[-1])
        print(f"{i+1:^20d} - {diff:>14.6e}, {f_value:>24.6e}")
        if diff < stop_eps:
            break
    # end solution
    return xs, ys
```

Using the functions implemented above, find the minimum of the above-defined target function ``f_quad``.
Use parameters values ``x_0 = 0, y_0 = 0`` for the starting point and a maximum of ``max_iter = 100`` iterations.
As stopping criteria, use ``stop_eps = 1e-6``.

Try to find the learning rate which requires a minimal amount of iterations to reach the function's minimum.

<!-- grade_id: cell-a77afc69c655, points: 1, grade: true, solution: true -->
```python
# Find a good learning rate!
# start solution
xs, ys = gradient_descent(0, 0, f_quad, f_quad_dw, f_quad_db, 0.1, 1e-6, 100)
# end solution
```

What happens if you choose learning rates which are significantly larger or smaller?

<!-- grade_id: cell-272f8398e36d, points: 1, grade: true, solution: true -->
```python
# Try out a significantly smaller learning rate and interpret the results!
# start solution
xs, ys = gradient_descent(0, 0, f_quad, f_quad_dw, f_quad_db, 0.001, 1e-6, 100)
# end solution
```

<!-- grade_id: cell-ed29b5af11ac, points: 1, grade: true, solution: true -->
```python
# Try out a significantly larger learning rate and interpret the results!
# start solution
xs, ys = gradient_descent(0, 0, f_quad, f_quad_dw, f_quad_db, 0.5, 1e-6, 10)
# end solution
```

The functions in the following code cell can be used to visualize the gradient descent process.
Make sure to execute the code cells to have the definitions loaded to your kernel.

```python
def draw_descent_steps(xs, ys):
    plt.plot(xs, ys, color=(1,1,1))
    plotted_steps = len(xs)
    plt.text(xs[0], ys[0], 1, color=(0,0,0))
    plt.text(xs[-1], ys[-1], plotted_steps, color=(0,0,0))
    
def plot_gradient_descent_process(xs, ys, target_fn, num_steps, x_true, y_true):
    if num_steps < len(xs):
        xs = xs[:num_steps]
        ys = ys[:num_steps]
    x_min = np.min(xs)
    x_max = np.max(xs)
    y_min = np.min(ys)
    y_max = np.max(ys)
    
    X, Y, Z = get_value_grid(x_min-1, x_max+1, y_min-1, y_max+1, target_fn)
    plot_2D_color_encoded(X, Y, Z)
    
    draw_descent_steps(xs, ys)
    
    plt.scatter(x_true, y_true, marker='o', color=(0,1,0))
    
```

Use the function ``plot_gradient_descent_process`` to visualize what happens in case of
* good choice of learning rate
* too small learning rate
* too large learning rate

<!-- grade_id: cell-254aeac7d9ef, points: 1, grade: true, solution: true -->
```python
# This cell: Good choice of learning rate!
# start solution
# 1) Use the function gradient_descent to perform gradient descent and get a sequence of x-y-Positions
xs, ys = gradient_descent(0, 0, f_quad, f_quad_dw, f_quad_db, 0.1, 1e-6, 10)
# 2) Use the function plot_gradient_descent_process to create a plot of the sequence
plot_gradient_descent_process(xs, ys, f_quad, 10, 3/2.5, -1/1.7)
# end solution
# Show plot
plt.show()
```

<!-- grade_id: cell-85fa8506ef38, points: 1, grade: true, solution: true -->
```python
# This cell: Too small choice of learning rate!
# start solution
# 1) Use the function gradient_descent to perform gradient descent and get a sequence of x-y-Positions
xs, ys = gradient_descent(0, 0, f_quad, f_quad_dw, f_quad_db, 0.001, 1e-6, 100)
# 2) Use the function plot_gradient_descent_process to create a plot of the sequence
plot_gradient_descent_process(xs, ys, f_quad, 100, 3/2.5, -1/1.7)
# end solution
# Show plot
plt.show()
```

<!-- grade_id: cell-164b2ef63afc, points: 1, grade: true, solution: true -->
```python
# This cell: Too large choice of learning rate!
# start solution
# 1) Use the function gradient_descent to perform gradient descent and get a sequence of x-y-Positions
xs, ys = gradient_descent(0, 0, f_quad, f_quad_dw, f_quad_db, 0.3, 1e-6, 10)
# 2) Use the function plot_gradient_descent_process to create a plot of the sequence
plot_gradient_descent_process(xs, ys, f_quad, 5, 3/2.5, -1/1.7)
# end solution
# Show plot
plt.show()
```