# Neural Network Basics - Exercises

## Exercise: Neuron computation

In the following code cell, the functions ``sigmoid()`` and ``compute_neuron()`` have already been set up.

Complete both functions to implement a) the sigmoid activation and b) the computations performed in a single artificial neuron, according to the lecture.

Execute the code cell to ensure your code is running. The expected result is ``0.631812``.

<!-- grade_id: cell-ae48b66a209b, points: 0, grade: false, solution: true -->
```python
import math
from typing import List


def sigmoid(input):
    ### YOUR SOLUTION HERE
    return 1/(1+math.e**(-input))
    ### END OF SOLUTION
    
def compute_neuron(weights: List[float], bias: float, inputs: List[float], activation_fn):
    weighted_sum = 0
    ### YOUR SOLUTION HERE
    for i in range(len(weights)):
        weighted_sum += weights[i] * inputs[i]
    weighted_sum += bias
    ### END OF SOLUTION
    activation = activation_fn(weighted_sum)
    return activation

weights = [1.0, 0.1, -0.3, 0.23, 0.51]
bias = -1.0
inputs = [1.0, 1.0, 1.0, 1.0, 1.0]

result = compute_neuron(weights, bias, inputs, sigmoid)
print(result)
```

Output:
```
0.6318124177361016
```

<!-- grade_id: cell-36d41813863b, points: 1, grade: true, solution: false -->
```python
weights = [1.0, 0.1, -0.3, 0.23, 0.51]
bias = 1.0
inputs = [1.0, 2.0, 0.5, -1.0, -0.5]

result = compute_neuron(weights, bias, inputs, sigmoid)
assert math.isclose(result, 0.8270696501038614, abs_tol=1e-9)
```

## Exercise: Layer computation

Complete the function ``compute_layer()`` to compute a whole layer of neurons.
The function parameters ``weight_list`` and ``bias_list`` contain one entry with weights and bias, respectively, for each neuron of the layer.

If implemented correctly, the result of the code cell should read ``[0.8668315902688739, 0.8192024316077874, 0.9386706246998785, 0.9585827129112119, 0.9672471083581451]``.

<!-- grade_id: cell-017f5498b4c3, points: 0, grade: false, solution: true -->
```python
import numpy as np
np.random.seed(89)

def compute_layer(weight_list, bias_list, inputs, activation_fn):
    neuron_values = []
    ### YOUR SOLUTION HERE
    for weights, bias in zip(weight_list, bias_list):   
        neuron_values.append(compute_neuron(weights, bias, inputs, activation_fn))
    ### END OF SOLUTION
    return neuron_values
        
        
weights_list = np.random.rand(5, 5)
bias_list = np.random.rand(5),
inputs = np.repeat(1.0, 5, 0)

results = compute_layer(weights_list, bias_list, inputs, sigmoid)
print(results)
```

Output:
```
[array([0.86683159, 0.88289367, 0.90110609, 0.93422161, 0.89824578])]
```

<!-- grade_id: cell-b613a940b6ac, points: 1, grade: true, solution: false -->
```python
np.random.seed(188)

weights_list = np.random.rand(5, 5)
bias_list = np.random.rand(5)
inputs = np.repeat(1.0, 5, 0)

results = compute_layer(weights_list, bias_list, inputs, sigmoid)
assert math.isclose(np.linalg.norm(results - np.array([0.9787163855567274,0.937879579058177,0.9852787697387955,0.9614546758005652,0.6570684566359081])),
                    0, abs_tol=1e-9)
```